# COMP SCI 1015 IAP - W09 - Practical

## Task 1: Counting words (30 marks)

Please write a function that tracks how many 3-letter words, 4-letter words, 5-letter words, and so forth are in a given paragraph.

Hint: use the `.split()` method to split a paragraph into a list of words.

```python
def word_length_count(para):
    cnt = {}
    li = para.split()

    # INSERT YOUR CODE BELOW
    # ~ 6 lines
    for word in li:
        length = len(word)
        if length in cnt:
            cnt[length] += 1
        else:
            cnt[length] = 1
    return cnt

# TEST CASE BELOW
para = 'It was the best of times, it was the worst of times'
cnt = word_length_count(para)
print(cnt) #{2: 4, 3: 4, 4: 1, 6: 1, 5: 2}   
    
```

Output:
```
{2: 4, 3: 4, 4: 1, 6: 1, 5: 2}
```

## Task 2: Rain fall (30 marks)

Below is a list of dictionaries which reports the amount of rainfall in major cities in the last week. Your task is to write codes that take a list of dictionary as input and output the accumulated rain fall for each city. 

```python
rains = [
    {
        'Sydney' : 50,
        'Melbourne' : 30,
        'Adelaide' : 22,        
    },
    {
        'Sydney' : 10,
        'Pearth' : 14
    },
    {
        'Adelaide' : 30,
        'Melbourne' : 50
    },
    {
        'Sydney' : 20,
        'Melbourne' : 30,
        'Brisbane' :50
    },
    {
        'Pearth' : 20,
        'Adelaide' : 30,
        'Brisbane' :50
    },
]

def accumulate_rain(rains):
    acc = {}
    
    # INSERT CODE BELOW
    # ~ 6 Lines
    for city_rain in rains:
        for city, rainfall in city_rain.items():
            if city in acc:
                acc[city] += rainfall
            else:
                acc[city] = rainfall
    return acc

# TEST CASE BELOW
r = accumulate_rain(rains)
print(r)
# {'Sydney': 80, 'Melbourne': 110, 'Adelaide': 82, 'Pearth': 34, 'Brisbane': 100}

```

Output:
```
{'Sydney': 80, 'Melbourne': 110, 'Adelaide': 82, 'Pearth': 34, 'Brisbane': 100}
```

## Task 3: Dictionary of tuple (40 marks)

The dicitonary `marks` contains the marks of students. Each student might have submitted multiple times, and the dictionary store all marks as a tuple. Please complete a function `highest_mark` that return a new dictionary that contians key-value pairs of each student and his/her highest mark.

If a student's max mark is over 100, you should set their mark to 100. If a student scored less than 0, set their mark to 0.

- hint1: You can assume all students submitted multiple times.
- hint2: consider writing a `find_max_mark` function


```python
import math

marks = {
    'a112':(40, 80, 70),
    'a113':(40, 80, 70, 40, 90, 70),
    'a114':(90, 28),
}

def find_max(marks_list):
    return max(min(mark, 100) for mark in marks_list)

def highest_mark(marks):
    dic = {}
    # INSERT YOUR CODE BELOW
    # write an additonal find_max function might make things easier
    # ~ 3 lines
    for student, marks_tuple in marks.items():
        max_mark = find_max(marks_tuple)
        dic[student] = max_mark
    return dic

...
    
highest_mark(marks) # {'a112': 80, 'a113': 90, 'a114': 90}
```

Output:
```
{'a112': 80, 'a113': 90, 'a114': 90}
```

Tests for Task1 (points: 30)
```
>>> word_length_count('uni of adelaide')
{3: 1, 2: 1, 8: 1}
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> word_length_count('good morning')
{4: 1, 7: 1}
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task2 (points: 30)
```
>>> r = [{'syd':25, 'ade':30}, {'mel':15}]
>>> accumulate_rain(r)
{'syd': 25, 'ade': 30, 'mel': 15}
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> r = [{'syd':25, 'ade':30}, {'syd':85, 'ade':10}]
>>> accumulate_rain(r)
{'syd': 110, 'ade': 40}
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task3 (points: 40)
```
>>> m = {'a':(30, 25), 'b':(98, 2), 'c':(36, 25)}
>>> highest_mark(m)
{'a': 30, 'b': 98, 'c': 36}
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> m = {'d':(90, 35), 'e':(68, 21), 'f':(72, 25)}
>>> highest_mark(m)
{'d': 90, 'e': 68, 'f': 72}
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> m = {'h':(0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0), 'i':(99, 110), 'j':(-1,-1)}
>>> highest_mark(m)
{'h': 0, 'i': 100, 'j': 0}
```
failure_message: Incorrect answer, try again!
success_message: Good job!
