# COMP SCI 1015 IAP - W09 - Workshop

## Demo 1 - Dictionary

A dictionary is like a list, but more general. When using a list, the indices must be integers, e.g. `li[3] = 5`. When using a dictionary, the indices can be different data types such as string, e.g. `dict['apple'] = 10`. This flexibility makes dictionary a convenient tool in many practical applciations.

A `list` contains a series of values, wheras a `dictionary` contains a collection of `keys` and `values`. Each `key` is associated with a single `value` and we call it a **key-value pair**. In the example above, `apple` is a key, and its corresponding value is `10`.

### Creating a dictionary 

We can create an empty dictionary using a curly bracket `{}` or the `dict` function. .

```python
# both statements create an empty dictionary
eng2jp = {} # method 1
eng2jp = dict() # method 2
```

### Add key-value pair

```python
eng2jp['bye'] = 'sayonara'
eng2jp['hi'] = 'konichiwa'
```

Codes above create a mapping from English word `bye` and `hi` to their respective Japanese translation of `sayonara` and `hi`. We can also initialise the dictionary as below

```python
# this is equivalent to the cell above.
eng2jp= {'bye':'Sayonara', 'hi':'konichiwa'}
```

Just like list, we can access the value using the square bracket operator `[]`. Just that the indices are now *strings*.

```python
print('The JP translation of "bye" is ' + eng2jp['bye'])
print('The JP translation of "bye" is ' + eng2jp['hi'])
```

### Modify key-value pairs 

Modifying values in a dictionary uses a similar syntax to modifying values in a list. 

```python
# create a dictionary with multiple entries
my_computer = {
    'CPU' : 'Intel I7',
    'Ram' : '16Gb',
    'OS' : 'Windows'    
}
print(my_computer)

my_computer['Ram'] = '8GB' # modify the key-value pair
print(my_computer)
```

We can also remove an item using `pop()`, just like list.

```python
my_computer.pop('CPU')
print(my_computer)
```

### Check if a key exists

When you try to access a key that is not existent in the dictionary, you will receive a `KeyError`.

```python
print(eng2jp['please'])
```

To avoid such error, you can first check if the key exists using the `in` keyword, before access its value.

```python
k = 'bye'
if k in eng2jp: # use `in` to check if the key exists
    print(f'<{k}> is in the dictionary')
else:
    print(f'<{k}> cannot be found')
```

---
## Activity 1 - Creating a codebook

<!-- BEGIN QUESTION -->

Caesar cipher is a type of substitution cipher in which each letter in the plaintext is replaced by a letter some fixed number of positions down the alphabet. For example, if the shift is `2`, then character `a` is replaced by `c`, and `b` is replaced by `d`. Please complete the following function that create a codebook with such mapping.

```python
alphabet = 'abcdefghijklmnopqrstuvwxyz'

def caesar_codebook(shift):
    codebook={}
    # INSERT YOUR CODE
    # ~ 4 lines
    length = len(alphabet)
    for i in range(length):
        original_char = alphabet[i]
        shifted_char = alphabet[(i + shift) % length]
        codebook[original_char] = shifted_char
    return codebook

def encrypt(msg, cb):
    s = ''
    # INSERT YOUR CODE
    # ~ 4 lines
    for char in msg:
        if char in cb:
            s += cb[char]
        else:
            s += char
    return s

# TEST CASES
cb = caesar_codebook(5)
s = encrypt('have a great day', cb)
print(s) # mfaj f lwjfy ifd
s = encrypt('adelaide', cb)
print(s) # fijqfnij
```

Output:
```
mfaj f lwjfy ifd
fijqfnij
```

<!-- END QUESTION -->

---
## Demo 2 - Dictionary as counters

Dictionaries are commly used as a collection of counters. For example, given a paragraph, we can use a dictionary to count how many times each letter or word appears. In the code below we can count how many times `COVID-19` and `symptoms` appear in the paragraph.

```python
paragraph = '''
If you have COVID-19 it can take several days to develop symptoms - but you’re contagious during this time. You are no longer contagious 10 days after your symptoms began.

The best way to avoid spreading COVID-19 to others is to:

Stay 6 feet away from others whenever possible.
Wear a cloth mask that covers your mouth and nose when around others.
Wash your hands often. If soap isn’t available, use a hand sanitizer that contains at least 60% alcohol.
Avoid crowded indoor spaces. Open windows to bring in outdoor air as much as possible.
Stay self-isolated at home if you are feeling ill with symptoms that could be COVID-19 or have a positive test for COVID-19.
Clean and disinfect frequently touched surfaces.
'''
# set up counters
cnt = {}
cnt['COVID-19'] = 0
cnt['symptoms'] = 0

for word in paragraph.split(): # split the paragraph into a list of words and loop through
    if word in cnt: # check if in COVID
        cnt[word] += 1
print(cnt)
```

---
## Activity 2 - Count characters 

<!-- BEGIN QUESTION -->

Similar to the demo above, please create a dictionary that counts the number of appearances for each alphabet.

```python
s = 'abcdefghijklmnopqrstuvwxyz'
cnts = {}

# STEP 1
# You should crate a dicitonary `cnts` that has 26 keys, from `a` to `z`, each key has a value `0`. 
# INSERT YOUR CODE BELOW
# Use for loop
# ~2 lines
for char in s:
    cnts[char] = 0
print(cnts['a']) # 0
print(cnts['d']) # 0

# STEP 2
# Given a sentence, loop through every character and calculate how many times each character appears.

words = 'an apple a day keeps the doctor away'
# INSERT YOUR CODE BELOW
# Use for loop
# ~3 lines
for char in words:
    if char.isalpha():
        char = char.lower()
        cnts[char] += 1
print(cnts['a']) # 6
print(cnts['d']) # 2
```

Output:
```
0
0
6
2
```

<!-- END QUESTION -->

---
## Demo 3 - loop through all key-value pairs

We can traverse all keys in a dictionary using a `for` loop (meaning, dictionary is iterable). Remember the pattern of for loop below, you will use it a lot.

```python
# create a dictionary
user_ids = {
    'Bob':102,
    'Chelsea':103,
    'Alice':101, # 'Alice' is a key connects to a value of 101
    'Dean':159,
    'Eagle':325
}

print(user_ids)
for k in user_ids:
    print(f'{k} - {user_ids[k]}')
```

---
## Activity 3 - Reverse loopup

<!-- BEGIN QUESTION -->

Given a dictionary `dict` and a key `k`, it is trivial to find the corresponding value `v = d[k]`. This operation is called **lookup**. 

Using dictionary, we can easily look up the corresponding value of a key `k`, such as `dict['k']`. However, what if we would like to do the opposite, meaning, given a value `v` and looks for its key? Unfortunately, there is no simple way to do so, and you would need to loop through the entire dictionary to achieve that.

Let's write such a lookup function below.

```python
dict1 = {
    'Alice':105,
    'Bob':112,
    'Chelsea':125,
    'Deb' : 134,
    'Edgard' : 107,
    'Fox' : 109
}

def reverse_lookup(di, v):
    # INSERT YOUR CODE
    # ~ 3 lines
    keys = [key for key, value in di.items() if value == v]
    if keys:
        return keys[0]
    else:
        return None
    
# === test case below ===
k = reverse_lookup(dict1, 134)
print(k) # Deb

k = reverse_lookup(dict1, 107)
print(k) # Edgard

k = reverse_lookup(dict1, 110)
print(k) # None
    
```

Output:
```
Deb
Edgard
None
```

<!-- END QUESTION -->

## Demo 4 - Tuples 

Now, let's switch to **tuple**, another built-in type in Python. A **tuple** is a sequence of values. Just like a **list**, the values can be any type, and they are indexed by integers. The main difference is that **tuples** are **immutable** (recall string).

We can create a tuple with a series of comma-separated vlaues or using the parentheses.

```python
t1 = ('a', 'b', 'c', 'd', 'e') # syntax 1
t2 = 'f', 'g', 'h', 'i', 'j' # syntax 2

print(f'{type(t1)} - {t1}')
print(f'{type(t2)} - {t2}')
```

Most list operators also work for tuples

```python
print(t1[2]) # access element
print(t1[:2]) # slices

for e in t1: # acces each element
    print(e)

```

You can also concat two tuples together using the `+` operator. Note that we did not modify `t1` or `t2`. Here we just create a new tuple `t3`.

```python
t3 = t1 + t2
print(t3)
```

### Tuple as return value 

<font color='red'> VERY IMPORTANT</font>

Tuple is especially useful when used as a return value. By definition, a function can only return one value. However, we can cirumvent this limitation using a tuple. For example, we can write a new `divide` function that returns both the quotient and remainder:

```python
def divide(dividend, divisor):
    q = int(dividend / divisor)
    r = dividend % divisor
    return q, r

t = divide(10, 7)
print(t)

q, r = divide(23, 3)
(q, r) = divide(23, 3) # line 9 and line 10 are equivalent.
print(f'quotient: {q} remainder: {r}')

```

---
## Submission Exercise

Please write a function `find_min_max` that takes a dictionary as an input and return a tuple of two keys to the min and max values respectively.

<!-- BEGIN QUESTION -->



```python
import math

def find_min_max(d):
    _min = math.inf
    _max = -math.inf
    
    k_min = 0
    k_max = 0

    # INSERT YOUR CODE BELOW
    # Loop through entire dictionary and compare min and max
    # ~ 6 Lines
    for key, value in d.items():
        if value < _min:
            _min = value
            k_min = key
        if value > _max:
            _max = value
            k_max = key
    return k_min, k_max
    
d1 = {'a':1, 'b':3, 'e':14}
(k1, k2) = find_min_max(d1)
print(f'{k1}, {k2}') # a, e

d1 = {'a':23, 'b':3, 'c':14, 'e':32, 'g':-14, 'kk':59, 'be':-32}
(k1, k2) = find_min_max(d1)
print(f'{k1}, {k2}') # be, kk

```

Output:
```
a, e
be, kk
```

<!-- END QUESTION -->

---
## Extention

### Tuple and dictionary 

As a bonus, Python dictionary has a very handy `.item` method that returns a sequence **key-value** tuples. 

```python
user_ids = {
    'Bob':102,
    'Chelsea':103,
    'Alice':101, # 'Alice' is a key connects to a value of 101
    'Dean':159,
    'Eagle':325
}

for (key, val) in user_ids.items():
    print(f'{key} - {val}')
```

### Dictionary and List

`List` can appear as values in the dictionary. For example, you might have a dictionary where each entry contains a list of English names beginning with each letter:
- A : ARIA, ARIBELLA, ARIS
- B : BARBARA, BELLA, BETSY
- C : CHLOE, CINDY, CLAUDIA

Now, a key value would correspond to a list of values. We can create the dictionary as following:

```python
names = {}
names['A'] = ['ARIA', 'ARIBELLA', 'ARIS']
names['B'] = ['BARBARA', 'BELLA', 'BETSY']
names['C'] = ['CHLOE', 'CINDY', 'CLAUDIA']
```

And we can loop through all names just like we did with nested list

```python
for key in names:
    for name in names[key]:
        print(name)

```

Tests for Submission (points: 0)
```
>>> d1 = {'a':8, 'b':3, 'e':2, 'f':13}
>>> find_min_max(d1)
('e', 'f')
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> d1 = {'1':89, '2':30, '3':12, '4':23}
>>> find_min_max(d1)
('3', '1')
```
failure_message: Incorrect answer, try again!
success_message: Good job!
