# COMP SCI 1015 IAP - W05 - Practical

## Task 1: Maximum and minimum (30 marks)

<!-- BEGIN QUESTION -->

Please complete the function below that return the element with minimum value in a given list. You need to complete the functions using three different loop techniques: `while`, `for`, and `for + range` respectively. 

Note that you will **NOT** get marks if you use the built-in Python function such as `sort()`, `min()`, or `max()`.

```python
import math

def min_for(li): # for loop
    # INSERT YOUR CODE BELOW
    # ~ 4 lines
    m = li[0]
    for num in li:
        if num < m:
            m = num
    return m

def min_while(li): # while loop
    # INSERT YOUR CODE BELOW
    # ~ 4 lines
    m = li[0]
    i = 0
    while(i<len(li)):
        if li[i] < m:
            m = li[i]
        i+=1
    return m
    
    
def min_range(li): # for + range
    # INSERT YOUR CODE BELOW
    # ~ 4 lines
    m = li[0]
    for x in range (1,len(li)):
        if li[x] < m:
            m = li[x]
    return m




# ================
# TEST CASES BELOW
def test(li):
    v1 = min_for(li)
    v2 = min_while(li)
    v3 = min_range(li)
    print(f'{v1} : {v2} : {v3}')

li = [3, 4, 2, 15, 17, 30, 5, 45]
test(li) # 2 : 2 : 2

li = [5, 8, 12, -5, -7]
test(li) #-7 : -7 : -7
```

Output:
```
2 : 2 : 2
-7 : -7 : -7
```

<!-- END QUESTION -->

## Task 2 - palindrome check (30 marks)

A **palindrome** is a word, phrase, number or sequence of words that reads the same backward as forward. Please follow the instructions below, in sequence, to complete your palindrome check function.

### Task 2.1 

Complete the function `reverse_string` function that takes one string and return its reverse.

```python
def reverse_string(s):
    ns=''    
    for c in s:
        # COMPLETE THE CODE BELOW
        # ~ 1 LINE        
        ns = c + ns
    return ns

# ================
# TEST CASES BELOW
s = reverse_string('test.')
print(s) # '.tset'
```

Output:
```
.tset
```

### Task 2.2

Please write a `compare_string` function that compare if two strings are identical. The spec of the function is
- it has two parameters `str1` and `str2`. You can assume `str1` and `str2` have same length.
- it returns `-1` if the content of str1 is identical to str2 (as a practice, please compare both string character by character.)
- it returns an integer `n`, where `n` is the index of the first character that is different between `str1` and `str2`

<font color='red'>Note, you should not use built-in string comparison methods. Instead, your code must iterate through each characters and compare</font>

```python
def compare_string(str1, str2):
    r = -1
    # INPUT YOUR CODE BELOW
    # Iterate through each characters and compare
    # ~ 6 LINES
    str1 = input("Please enter your string 1: ")
    str2 = input("Please enter your string 2: ")

    for i in range(len(str1)):
        if str1[i] != str2[i]:
            return i
    return r
        
# ================
# TEST CASES BELOW
a = compare_string('test', 'test')
print(a) # -1
a = compare_string('xest', 'test') # first char is different
print(a) # 0
a = compare_string('text', 'test') # 3rd char is different
print(a) # 2
```

Output:
```
0
-1
2
```

### Task 2.3 

Please write a **is_palindrome** function that
- takes one string parameter
- return True if the string is a palindrome and return False otherwise

```python
def is_palindrome(s):
    # INSERT YOUR CODE BELOW
    # USE THE ABOVE TWO FUNCTIONS ABOVE TO COMPLETE THE FUNCTION
    # ~ 6 LINES
    s = ''.join(c for c in s.lower() if c.isalnum())
    return s == s[::-1]
    
p = is_palindrome('anna') # True
print(p)
p = is_palindrome('civic') # True
print(p)
p = is_palindrome('alice') # False
print(p)
```

Output:
```
True
True
False
```

## Task 3 : Caesar cipher (40 marks)

The Caesar cipher is one of the earliest known and simplest ciphers. It is a type of substitution cipher in which each letter in the plaintext is 'shifted' a certain number of places down the alphabet. For example, with a shift of 2, the word "hi", will become "jk". Because **h** is replaced by **j**, **i** is replaced by **k**.

In this task, you will write a program that encrypt the user input with Caesar Cipher, and decrypt it. For simplicity, we can assume that the user will only use lowercase letters. 

Below are some helper code for you. the 
- char_list containing all 26 lowercase character, 
- the `.index` function returns the index of a value in the list.

```python
# HELPER CODE
char_list = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j', 'k', 'l', 'm', 'n', 'o', 'p', 'q', 'r', 's', 't', 'u', 'v', 'w', 'x', 'y', 'z']

# below is a quick example of shifting the character by 2
idx = char_list.index('h')          #get the index of character 'h'
idx2 = (idx + 2 ) % len(char_list)  #shift the index by 2
new_char = char_list[idx2]          #get the encrypted character using the new index

print(f'h -> {new_char}')
```

### Task 3.1 Encrypt

Once you understand the code below, please complete the `caesar_encrypt` function below, the function has two parameters
- the input string. Note that the input string might contain space character
- `n`, the number of shift 

```python
def caesar_encrypt(s, n):
    enc = '' # the encrypted string

    # COMPLETE THE CODE BELOW
    # iterate through each character and shift it by n, you should skip the space character
    # ~ 7 lines
    for char in s:
        if char.isalpha():
            shifted = (ord(char.lower()) - ord('a') + n) % 26 + ord('a')
            enc += chr(shifted)
        else:
            enc += char
    return enc

# TEST CASES BELOW
es = caesar_encrypt('an apple a day', 4)
print(es) # er ettpi e hec

es = caesar_encrypt('keeps the doctor away', 8)
print(es) # smmxa bpm lwkbwz ieig
```

Output:
```
er ettpi e hec
smmxa bpm lwkbwz ieig
```

### Task 3.2 Decrypt
Following Task 1.1, when the receiver received the encrypted message, a decryptor is needed to translate the encrypted word into a plain one. Please write a program that decrypt the encrypted message.


```python
def caesar_decrypt(s, n):
    # COMPLETE THIS FUNCTION
    # HINT: Use the caesar_encrypt but with a different shift :)
    # ~1 LINE 
    result = ''
    for char in s:
        if char.isalpha():
            shifted = (ord(char.lower()) - ord('a') - n) % 26 + ord('a')
            result += chr(shifted)
        else:
            result += char
    return result

# TEST CASE BELOW
n = 12
s = 'an apple a day keeps the doctor away'
es = caesar_encrypt(s, n)
print(es) # mz mbbxq m pmk wqqbe ftq paofad mimk
ds = caesar_decrypt(es, n)
print(ds) # an apple a day keeps the doctor away
```

Output:
```
mz mbbxq m pmk wqqbe ftq paofad mimk
an apple a day keeps the doctor away
```

---

Tests for Task 3.1 (points: 20)
```
>>> caesar_encrypt('this is a testing string', 5) == 'ymnx nx f yjxynsl xywnsl'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> caesar_encrypt('adelaide', 23) == "xabixfab"
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task1 (points: 30)


Tests for Task2.1 (points: 10)
```
>>> reverse_string('an apple') == 'elppa na'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> reverse_string('doctor away') == 'yawa rotcod'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task2.2 (points: 10)
```
>>> compare_string('happy', 'happy') == -1
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> compare_string('habby', 'happy') == 2
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> compare_string('hapBy', 'happy') == 3
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task2.3 (points: 10)
```
>>> is_palindrome('apple') == False
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> is_palindrome('civic') == True
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> is_palindrome('kayak') == True
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task3.2 (points: 20)
```
>>> caesar_decrypt('xabixfab', 23) == 'adelaide'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> caesar_decrypt('ixxtm', 8) == 'apple'
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
