```python
# Initialize Otter
import otter
grader = otter.Notebook("2023_W01_Workshop-Submission.ipynb")
```

# COMP SCI 1015 IAP - W01 - Workshop submission

## Submission Exercise
This is your first workshop submission. Please write the code that print out `Hello World!`

<!-- BEGIN QUESTION -->



```python
# INSERT YOUR CODE BELOW
print("Hellow World!")
```

Output:
```
Hellow World!
```

<!-- END QUESTION -->

## <font color='red'>Remember to submit your work to Gradescope and sign out with yout tutor before leaving the session.</font>

Tests for Task1 (points: 100)
```
>>> out = get_print_output(); # retrieves the print output.
>>> reset_print_output(); # Resets the print output.
>>> 'Hello World' in out
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
