# COMP SCI 1015 IAP - W10 - Practical

# Task 1: list substraction (60 marks) 

Please write two functions that subtract one list from another. 
- `list_sub_v1(li1, li2)` directly remove elements in `li2` from `li1`
- `list_sub_v2(li1, li2)` create a new list that contains the result of `li2` - `li1`. For example, `[1, 2, 2, 3] - [2, 3] = [1, 2]`

Please write one function that shows the different elements between two lists
- `diff(li1, li2)` create a new list that contains all elements that are in `li1` but not `li2`. Please try completing this function using `list_comprehension` (you will still get mark if **not** using list comprehenship)


```python
def list_sub_v1(li1, li2):
    # INSERT YOUR CODE BELOW
    # ~ 3 LINES
    for item in li2:
        if item in li1:
            li1.remove(item)
    
def list_sub_v2(li1, li2):
    li3 = []    
    # INSERT YOUR CODE BELOW
    # ~ 6 LINES
    li3 = []
    for item in li1:
        if item not in li2:
            li3.append(item)
    # END OF YOUR CODE
    return li3
                
def diff(li1, li2):
    # INSERT YOUR CODE BELOW
    # 1 line if using list comprehension
    return [item for item in li1 if item not in li2]

# ================
# TEST CASES BELOW
li1 = [1, 2, 3, 4]
li2 = [2, 3]
list_sub_v1(li1, li2)
print(f'1 - {li1}') # 1 - [1, 4]

li1 = [1, 2, 3, 4]
li2 = [2, 3]
li3 = list_sub_v2(li1, li2)
print(f'2 - {li1} {li3}') # 2 - [1, 2, 3, 4] [1, 4]

li1 = [1, 2, 2, 3, 3, 4]
li2 = [2, 3]
list_sub_v1(li1, li2) # note that we remove only one '2' from li1
print(f'3 - {li1}') # 3 - [1, 2, 3, 4] 

li1 = [1, 2, 2, 3, 3, 4]
li2 = [2, 3]
li3 = list_sub_v2(li1, li2)
print(f'4 - {li1} {li3}') # 4 - [1, 2, 2, 3, 3, 4] [1, 2, 3, 4]

li1 = [1, 2, 2, 3, 3, 4]
li2 = [2, 3]
li3 = diff(li1, li2)
print(f'5 - {li3}') # 5 - [1, 4]
```

## Task 2 - dice game (40 marks) 

In this task, you will create a dice rolling game. In each round, each player rolls all the dice. The player with the highest sums win the round. Please read the code and comments below and complete the functions:
- `add_player`: add a new player into the game. It will only be called before the game starts.
- `roll`: One round of the dice game. All players roll all the dice and the one with max score win the round.
- `display`: Display the history of each round

The `game['round']` stores the score for each players in each round. For example, if we have 3 players A, B, C
- Round 1, A rolls 15, B rolls 12, C rolls 5
- Round 2, A rolls 4, B rolls 9, C rolls 15

Then the content of the game will be:
- `game['num_players']` is `2`
- `game['num_rounds']` is `2`
- `game['players']` is `['A', 'B', 'C']`
- `game['winners']` is `['A', 'C']`
- `game['round']` will be `[[15, 12, 5],[4, 9, 15]]`

```python
import random

game = {
    'num_dice' : 3, # number of dice. A constant.
    'num_players' : 0, # number of players
    'num_rounds' : 0, # number of rounds played
    'players' : [], # player names in a list
    'winners' : [], # winners for each round in a list
    'round' : [] 
    # nested list storing the score for each players in each round
}

def add_player(game, name):
    game['players'].append(name) # add player name
    game['num_players'] += 1 # increase player number

def roll(game):
    # INSERT YOUR CODE BELOW
    # Play a round of the game 
    # ~ 20 lines
    round_scores = []
    for _ in range(game['num_players']):
        score = sum([random.randint(1, 6) for _ in range(game['num_dice'])])
        round_scores.append(score)
    max_score = max(round_scores)
    winners = [game['players'][i] for i, score in enumerate(round_scores) if score == max_score]
    game['winners'].extend(winners)
    game['round'].append(round_scores)
    game['num_rounds'] += 1
    
def display(game):
    # DON'T CHANGE THE CODE IN THIS FUNCTION
    s = ''
    for r in range(game['num_rounds']):
        print(f'------\nRound {r}:')
        for i in range(game['num_players']):
            print(f"{game['players'][i]} {game['round'][r][i]}")
        print(f"winner: {game['winners'][r]}")
    return s
        
#================
#TEST CASE BELOW
random.seed(33) # comment this line if you would like to see random results
add_player(game, 'Alice')
add_player(game, 'Bob')
add_player(game, 'Christine')
roll(game)
roll(game)
roll(game)
display(game)

# THE OUTPUT SHOULD BE 
# ------
# Round 0:
# Alice 13 
# Bob 9 
# Christine 16 
# winner: Christine
# ------
# Round 1:
# Alice 13 
# Bob 14 
# Christine 13 
# winner: Bob
# ------
# Round 2:
# Alice 15 
# Bob 8 
# Christine 12 
# winner: Alice
```

Output:
```
------
Round 0:

Alice 13
Bob 9
Christine 16
winner: Christine

------
Round 1:

Alice 13
Bob 14
Christine 13
winner: Bob

------
Round 2:

Alice 15
Bob 8
Christine 12
winner: Alice

''
```

Tests for Task1 (points: 60)
```
>>> li = [5, 6, 7, 8]
>>> list_sub_v1(li , [6, 8])
>>> li
[5, 7]
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> li3 = list_sub_v2([12, 13, 15], [12])
>>> li3
[13, 15]
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> li3 = diff([23, 56, 78], [78])
>>> li3
[23, 56]
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task2 (points: 40)
```
>>> random.seed(15) # comment this line if you would like to see random results
>>> game = {'num_dice' : 3,'num_players' : 0, 'num_rounds' : 0,'players' : [], 'winners' : [], 'round' : [] }
>>> add_player(game, 'D')
>>> add_player(game, 'E')
>>> add_player(game, 'F')
>>> roll(game)
>>> roll(game)
>>> s = display(game)
>>> s
'------\nRound 0:\nD 8 \nE 9 \nF 4 \nwinner: E------\nRound 1:\nD 14 \nE 6 \nF 13 \nwinner: D'
```
failure_message: Incorrect answer, try again!
success_message: Good job!
