# COMP SCI 1015 IAP - W03 - Workshop 

## Demo 1: Funciton

### Built-in functions

A `function` is a block of code that is designed to do **one** specific job. For example, we have been using the Python built-in **print()** function for a couple weeks now, and its job is to print out text on the display. We have also used the **int()** function to perform the *type casting*. 

```python
print('hello')
```

```python
int('35')
```

We have also used some math functions, such as `math.sqrt()` last week. Here `math` is the **module**, and `sqrt` is a function defined in that `module`. You can find the complete list of functions inside the math module at https://docs.python.org/3/library/math.html. By importing different **modules**, you will have access to all kind of different utility functions built by programmers around the globe!

```python
import math

print(math.pow(4, 2)) # 4 square
print(math.log2(4)) # 
```

Some modules also define **constants** for easier use. For example:

```python
print(math.pi)
print(math.e)
```

### Function composition 

The functions become even more powerful when you **compose** them together. For example:

```python
print(math.log2(math.pow(4, 2)))
print(math.sin(math.pi))
print(math.sin(math.pi*0.5))
```

### Defining your own functions
In Python, the naming rule of the function is the same as a variable. It starts with a character and the name can contains both character, underscore, and digits. Below we define a function `hello_world`:

```python
def hello_world(): # function definition
    print('Hello')
    print('World!')
```

The first line is the *header* while the rest is the *body*. Note that every lines in the body must be indented. By convention, indentations are always four spaces and Python uses this indentation to decide which part of the code is the function body.

After defining the function, we can use the function with a *function call*.

```python
hello_world() # note the parenthesis 
```

Similarly, you can also define multiple different functions and put them together. Note the **flow of execution** here too! Python first execute `hello_3`, then in turn executes 3 `hello`.

```python
def hello():
    print("hello")
    
def hello_3():
    # call function hello 3 times
    hello() 
    hello()
    hello()

hello_3()

```

---
## Activity 1 

The code below return a random floating-point number. Please write a function that prints a random **integer** number between 0 to 100. 

```python
import random
random.random() # <-- this line give you a random number between 0 and 1, try print it

def random_int():
    # YOUR CODE BELOW
    # ~ 1 LINE
    r = random.randint(0,100)
    print(r)

random_int()
```

Output:
```
88
```

---
## Demo 2: Argument and Parameter

### Argument 

Some functions require *arguments*. For example, when you call the `print` function, you need to supply a string. Some functions require more than one argument, like the `math.pow` function that calculate the power of power of a base number

```python
import math

print('hello world') # 'hello world' is an argument to the print function

math.sqrt(30)        # 30 is an argument to the math.sin function 

math.pow(2, 5)       # 2, 5 are arguments to the math.pow function
```

### Parameter 

Inside the function, `arguments` are assigned to variables called `parameters`. For example, here we define a function `say_hello` with a single parameter `name`. 

```python
def say_hello(name):
    print(f'Hello {name}!')
```

Then we call the `say_hello` function providing an argument `Alice`. You can see that the value `Alice` is assigned to the variable `name` and get printed!

```python
say_hello('Alice')
```

---
## Activity 2 - Add and Multiplication

In this activity, please write your own math functions:
- an **add** function that takes two parameters x and y. The function prints out the value of x + y
- a **mul** function that takes tree parameters x, y, and z. The function prints out the value of x * y * z

```python
def add(x, y):
    # INSERT YOUR CODE BELOW
    # 1 LINE
    s = x + y
    print(s)
    
def mul(x, y, z):
    # INSERT YOUR CODE BELOW
    # 1 LINE    
    m = x * y * z
    print(m)
    
add(5, 10) # print 15
mul(2, 3, 4) # print 24

```

Output:
```
15
24
```

---
## Demo 3: Scope of variables

In programming languages, **scope** defines where you can access a variable. Some variables can be accessed everywhere in the code and we say these variables have **global** scope. Whereas other variables, such as defined inside functions, have a **local** scope. See examples below:

```python
def func1():
    x = 3 # x has a local scope and can only be referenced inside the function

func1()
print(x) # Python will report an error! Because variable x is local and only existed inside the function
```

```python
y = 5 # y is defined outside the function, and is 'global'

def func2():
    print(y) # thus y can be accessed inside the function
    
func2()

```

However, if you execute the code below, you might be surprised about the result.

```python
y = 5 

def func3():
    y = 1 # <- Python create a NEW local variable y in the function
    
func3()
print(y) # You might think it will print 1

```

It is because in line 4, Python create a new variable `y` in the **local** scope of `func3` and update its value. Although the variable share the same symbol `y`, it has no effect to the global variable `y` declared in line 1.

To update the value of a global variable inside a function, you would need to use the `global` keyword to inform Python that you intend to reference and update a global variable inside the function.

```python
y = 5 

def func3():
    global y # Tell Python to reference the global variable y
    y = 1 
    
func3()
print(y)

```

---
## Activity 3.1 - Counter

Please update the code below so that your program counts how many times the function `hello` is called.

Hint: Notice `cnt` is defined outside of the function, how do you change the variable from inside the function? Review the lecture demo if you are struggling!

```python
cnt = 0
def hello():
    # INSERT YOUR CODE BELOW
    # 2 line
    global cnt
    cnt = cnt + 1
    
    
hello()
hello()
hello()
print(cnt) # should output 3
```

Output:
```
3
```

---
## Activity 3.2 - Mean

Please write two functions that calculate *mean*.
- mean_g() calcualtes the mean of global variables x, y, z
- mean_l(a, b, c) calcualtes the mean of local variables a, b, c
- mean_c(x, y, z) calcualtes the mean of local variables x, y, z

Also ask your self, will calling `mean_c` changes the value of the global variablex `x`, `y`, `z`? 

```python
x = 5
y = 6
z = 7

def mean_g():
    # 1 line
    m = (x + y + z)/3
    print(m)
    
def mean_l(a, b, c):
    # 1 line
    m = (a + b + c)/3
    print(m)
    
def mean_c(x, y):
    # 1 line
    m = (x + y + z)/3
    print(m)
    
mean_g() # 6
mean_l(5, 6, 7) # 6
mean_c(2, 3) # 4
```

Output:
```
6.0
6.0
4.0
```

---
## Demo 4: Functions with return value

Python functions can produce a `return value`. For example, many mathematic functions, like `math.sqrt`, provide a `return value`. In the code below, the `sqrt` function takes an argument of `25`, calculate its sqaure root, and return the value `5`. 

```python
import math

s = math.sqrt(25)
print(s)
```

So far, the functions we defined on our own are all `void`. Meaning, these functions perform some actions, like printing out some values, but  they don't return a value.

Below is a function with a `return value`. Can you figure out what the function `add` does?

```python
def add(x, y):
    c = x + y # c is a temporary variable that is only visible INSIDE the function.
    return c  # the return statement

s = 3 + 5
print(s)

s = add(3, 5) # add operation in a function form. The return value is assigned to the variable s.
print(s)
```

We can also simplify the function `add` a bit using the composition:

```python
def add(x, y):
    return x + y # composition
```

The function above looks cleaner. But haivng a temporary variable, like `c` above, would help you debug more easily. I encourage you to prioritize readability of your code over the length of your code.

---
## Activity 4

Please create the following functions:
- add(x, y), which outputs x + y
- mul(x, y), which outputs x * y

```python
# DEFINE ADD
# ~2 lines
def add(x,y):
    a = x + y
    return a

#DEFINE MUL
# ~2 lines
def mul(x,y):
    m = x * y
    return m

z = add(3, 5) + mul(2, 4) # (3+5) + (2*4)
print(z) # should print 16
```

Output:
```
16
```

---
## Submission Exercise

<!-- BEGIN QUESTION -->

Use the two functions `add` and `mul` above to perform the calculation below:

$(3+5)\times(6+7)\times(8+9)$

Please do the following:

1. calculate $3 + 5$, $6 + 7$, and $8+9$ respetively and assigned them to three variables `a1`, `a2`, `a3`
2. calculate $a1 \times a2$, and assigned the answer to $b$
3. calculates $b \times a3$

```python
# STEP 1add()
a1 = add(3,5)
a2 = add(6,7)
a3 = add(8,9)

# STEP 2
b = mul(a1,a2)

# STEP 3
ans = mul(b,a3)
print(ans) # 1768
```

Output:
```
1768
```

<!-- END QUESTION -->

## <font color='red'>Remember to submit your work to Gradescope and sign out with yout tutor before leaving the session.</font>

Tests for Submission (points: 0)
```
>>> assert ans == 1768
```
failure_message: Incorrect answer, try again!
success_message: Good job!
