# COMP SCI 1015 IAP - W08 - Workshop 

## Demo 1

### Writing a file 

A text file contians a sequence of characters. A text file could lie on our hard disk drive and we can read open the file and read its content using different applications, such as notepad on Windows.

Python also provides functions for reading and writing files. To write a file, we open the file with the `w` mode. 

```python
fout = open('output.txt', 'w')
```

Note that if the file `output.txt` already exists, Python will clear all the content in the file! So, be careful! Here the variable `fout` is a `file object` that represents the file. We can now use the `write` method to put data into it. 

Note that we need to call the `close()` method of the file object to make sure all the data are put into the file.

```python
line = 'hello world\n'
fout.write(line) # actually write the data
fout.close() # close the file object and flush the buffer (watch video)
```

So far, we have created a file `output.txt` and write a line of text into it. Now, where is this data? It will be at the same directory where you start your Jupyter notebook. 

If you are not sure which directory you are in, you can use the `os.getcowd()` method and obtain your current working directory.

```python
import os
cwd = os.getcwd() # get current woring directory
print(cwd)
```

Output:
```
/Users/benq0630/Downloads
```

If you would like to add more text into an existing text file, you can open the file in `append` mode.

```python
f = open('output.txt', 'a') # append mode
f.write('2nd line\n')
f.write('3rd line\n')
f.write('4th line\n')
f.close()
```

### Read a file

To read a file, we have to create the `file object` in the read mode using the argument `r`. The `read()` method will get all the content in the file at once. 

```python
fin = open('output.txt', 'r') # open the file in read mode
content = fin.read() # read all content
print(content)
fin.close() # close the file object
```

Output:
```
hello world
2nd line
3rd line
4th line
```

### `with` keyword 

Python also provides a covenient `with` keyword that handle the closing of file automatically. You are encourged to use this method to read / write your file. Try running the two cells below and check if the result is as you expected!

```python
with open('output.txt', 'w') as fout:
    fout.write('1st line\n')
    fout.write('2nd line\n')
    fout.write('3rd line\n')
    fout.write('4th line\n')
    # note that we do not need to close the file now.
```

```python
with open('output.txt', 'r') as fin:
    content = fin.read()
    print(content)
    # note that we do not need to close the file now.
```

Output:
```
1st line
2nd line
3rd line
4th line
```

---
## Activity 1: Read and write 

<!-- BEGIN QUESTION -->

Complete the function `read_and_write(src, dst)` below. The function reads the file content from a file at `src` and write it into a file at `dst`.

```python
def read_and_write(src, dst):
    # INPUT YOUR CODE BELOW
    # ~ 4 lines
    with open(src, 'r') as f1:
        content = f1.read()
    with open(dst, 'w') as f2:
        f2.write(content)
    
    
    
read_and_write('output.txt', 'test1.txt')
read_and_write('test1.txt', 'test2.txt')
# Check if you have both test1.txt and test2.txt in your folder
```

<!-- END QUESTION -->

---
## Demo 2 

So far we use the `read()` method to read the entire content in a file. Doing so is usually fine for smaller text file. However, when dealing a much larger file with tens of thousands of lines, reading the content line by line would be more efficient. We can do this by using the `readline()` method as below

```python
with open('output.txt', 'r') as fin:
    l = fin.readline()
    print(l)
    l = fin.readline()
    print(l)
```

When we read the file sequentially using `readline()` method, the file object needs to record how far we have read in the file. This is called an `offset`. We often read a large file sequentially. Thus the file object maintain an **offset** that record how far we have read in the file. 

The `tell()` method prints out the current offset and we can manually set the offset using the `seek()` method. For example, if we would like to re-read the whole file, we can reset the offset with `seek(0)`.

```python
with open('output.txt', 'r') as fin:
    offset = fin.tell()
    print(f'offset at the beginning: {offset}')
    line = fin.readline() # read first line
    line = fin.readline() # read second line
    line = fin.readline() # read third line
    print(line)
    offset = fin.tell()
    print(f'offset after reading three lines: {offset}')
    
    print('---reset the offset---')
    fin.seek(0) # reset
    print(f'offset is {fin.tell()}')
    line = fin.readline() 
    print(line)
    
    print('---move the offset to the end of 3rd line---')
    fin.seek(offset) # move to 3rd line
    print(f'offset is {fin.tell()}')
    line = fin.readline() 
    print(line)
```

### Read through a file line by line 

In Python, a `file object` is iterable. So we can directly use the `for` loop to go through the entire file.

```python
with open('output.txt', 'r') as fin:
    for line in fin:
        print(line)
```

---
## Activity 2: Read content into a list 

<!-- BEGIN QUESTION -->

The code below create a file containing 30 lines, where each line contains an integer. Pleae complete the function `read_to_list(src)` that read the content of the file and store them into a list.

```python
def read_to_list(src):
    # INSERT YOUR CODE
    # ~ 5 lines, should return a list
    arr = []
    with open(src, 'r') as fout1:
        for line in fout1:
            arr.append(int(line))
    return arr


# TEST CASE BELOW
with open('array.txt', 'w') as fout:
    for i in range(10):
        fout.write(f'{i}\n')
li = read_to_list('array.txt')
print(li) # should print [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]

```

Output:
```
[0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
```

<!-- END QUESTION -->

---
## Demo 3: Error vs Exception

All previous codes assume that the file actually exists. However, that might not always be the case. If you attempt to read a non-existent file, you will get an `FileNotFoundError`. 

Note that although Python reports a `FileNotFoundError`, it is different from the `syntax error`. Officially, the errors that are detected during the **execution** is called an **exeception**.


```python
f = open('test.xyz')
```

For example, the code below has an additional `)` and will generate a `Syntax Error`. It means that Python does not understand the code (cannot parse it) and cannot exsecute it at all.

```python
print('hello'))
```

In contrast, the code above generate an `exception`. Meaning the code has correct syntax and can be executed. However, an error was detected during the execution. In this case, the code tries to divide a number by zero, which is an undefined behaviour. This is an `exception` and Python raises a `ZeroDivisionError`.

```python
print(20/0)
```

Below is anoter example where the code attempts to cast a string into an integer, but Python detects an exception because it does not know how to convert a string consists of `abc` string to a number.

```python
a = '13'
b = int(a) # it is OK

a = 'abc'
c = int(a) # Python generats an exception 
```

Python provides a convenient way to handle such exception - the **try and catch blocks**. When Python encountered an exception in the **try** block, instead of quitting the program, it will instead execute the codes in the **except** block. You can try changing the value of **a** below and check the result.

```python
try:
    a = 'abc'
    #a = '15'
    c = int(a)
    print(f'successfully convert {a} to int')
except:
    print(f'cannot convert {a} to int')
```

Another useful exception is the index out of range.

```python
l = [1, 2, 3]
a = l[3]    
```

```python
try:
    a = l[3]
except:
    print("something wrong on the indexing")

```

And we can also handle the file non-existent error here.

**[Read more about exception and handle here](https://realpython.com/python-exceptions/).**

---
## Activity 3: Interactive adder with exception handling 

<!-- BEGIN QUESTION -->

Write a simple program that takes user inputs and add them together. 
- Your code should keep accepting user inputs until the user inputs 'quit'. 
- Please use try and except blocks to ensure the user input an integer

For example:

Please input the first number: 13
Please input the second number: 30
>>The answer is 43

Please input the first number: ab
>> please input integer only
Please input the first number: 5
Please input the second number: 8
>>The answer is 13

Please input the first number: quit
>> Bye!


```python
i = 0
while True:
    # 3 lines below prevents autograder from going into infinite loop
    # please do not change
    i = i + 1
    if i > 20:
        break
        
    # INSERT YOUR CODE BELOW
    # ~ 10 LINES
    try:
        fir_num = input("Please input the first number: ")
        if fir_num == "quit":
            print("Bye!")
            break
        fir_num = int(fir_num)
            
        sec_num = input("Please input the second number: ")
        sec_num = int(sec_num)
        ans = fir_num + sec_num
        print(f"The answer is {ans}")
    except ValueError:
        print("Please input integers only for first and second numbers.")


```

Output:
```
The answer is 30
The answer is 30
```

<!-- END QUESTION -->

---
## Demo 4: Raising Exceptions

It is also possible to create your own Exceptions. To do this we use the `raise` keyword.

*Note: Sometimes the action of "raising an exception" is called "throwing an exception" as some languages use the throw keyword instead, people will understand what you mean regardless of the phrasing you use.*

```python
x = input("Enter your name:") # prompt user for their name

if (x == ""):
    raise Exception("You cannot enter an empty string!")
```

As you have seen earlier, there are different types of exceptions, such as `TypeError` and `ValueError`. We can build our own Exceptions off these as well as shown below.

First, we could have our error classified as a built-in error while also giving new information

```python
x = 0

# check if x is a string, if not raise TypeError.
print("x's type is " + str(type(x)))
if type(x) != str:
  raise TypeError("Only integers are allowed") 
```

### Why would I raise my own Exceptions and how do I pass an exception around to different parts of my code to handle it appropriately?

Raising your own exceptions are especially valuable in cases where you might need to pass an Exception up several functions or in cases you just want to give your own more useful error messages. To achieve this, we can capture the exception using the `as` keyword as shown below.

```python
def verify_name(name):
    if (name == ""):
        raise Exception("Error code 001: No name provided.")
    elif ('-' in name):
        raise Exception("Error code 002: You cannot enter a hyphenated name. Please call HR on (08) 5550 XXXX for assistance.")
    else:
        return True

def get_name():
    name = input("Enter your name:")
    try:
        verify_name(name)
        return name
    except Exception as e:
        raise e

try:
    print("Hello " + get_name())
except Exception as e:
    print(str(e))
    raise e
```

In the above example, the control flow of the program was:
1. call the `get_name` function
1. get name prompts the user for input
1. the user input is checked in `verify_name`
   1. if the name is not acceptable, an error is passed onto `get_name`. At this time the tracelog is not yet dumped to the screen.
   1. `get_name` passes the error onto the first executed line (within the try). Since the try encountered an exception, control is passed onto the except block.
   1. the except block prints the error message
   1. as there is no further layers, the traceback is finally dumped to the screen.

<!-- BEGIN QUESTION -->

---
## Submission Exercise: Reverse text

Please createa a Python function **reverse_text(src, dst)** that read the file `src` and create another file `dst` with a line-by-line reverse content of file `src`. 

```python
def reverse_text(src, dst):
    with open(src, 'r') as fin:
        # INSERT YOUR CODE
        # Read each line from src and store them in reverse order into a list
        # ~ 3 lines
        lines = fin.readlines()
        lines.reverse()
        
    with open(dst, 'w') as fout:
        # INSERT YOUR CODE
        # Write each line into dst
        # ~ 2 lines
        for line in lines:
            fout.write(line)
            
# TEST CASE
with open('output.txt', 'w') as fout:
    for i in range(5):
        fout.write(f'{i}\n')

reverse_text('output.txt', 'reverse.txt')
# reverse.txt should have the following content (without #)
# 4
# 3
# 2
# 1
# 0

```

<!-- END QUESTION -->



Tests for Submission (points: 0)
```
>>> reverse_text('W08_test.txt', 'test.txt')
>>> with open('W08_test.txt') as fin:
...     t = fin.read()
...     print(t.strip(' \n'))
0
1
2
3
4
```
failure_message: Incorrect answer, try again!
success_message: Good job!
