# COMP SCI 1015 IAP - W08 - Practical

## Task 1: A simple e-book reader (50 marks)

When dealing with the input / output of large files, we usually do not need to hold all contents in the system memory. For example, if we are reading a very long book, we only need to keep the pages the user is currently reading in the memory. In this example, you will create a simple e-book reader that does just that.

Please download the `W08_LOTR.txt` file from MyUni. It is a big text file that contains the first volume of The Lord of The Ring. Your e-book reader function should take the input of a page number, then output the text of that particular page.

Some notes:

- Open W08_LOTR.txt with your text editor and figure out the format of page numbers. **There are some missing page number, and you can ignore those**.
- The file is big. You should sequentually read the file using the `readline()`, instead of using the `read()`.
- use `strip()` method to remove leading and trailing white spaces for each line.
- Your e-book reader function should take the input of a page number, then output the text of that particular page. You do not need to continuously take user input. 

```python
def read_page(page):
    with open('W08_LOTR.txt', 'r', encoding='utf8') as fin:
        flag = False
        next_page = page + 1
        text = ''
        # INSERT YOUR CODE BELOW
        # Hint:
        # Use the flag variable to control which lines to be added to s
        # Change the flag to true when we find "page"
        # Change the flag to false when we find "next_page"
        # ~ 8 lines
        for line in fin:
            if 'Page ' + str(page) in line:
                flag = True
            elif 'Page ' + str(next_page) in line:
                break
            if flag:
                text += line.strip() + '\n'
        return text
# TEST CASE BELOW
text = read_page(1)
print(text)

text = read_page(10)
print(text)
```

## Task 2: Codebook (50 marks)

In crypotography, a **codebook** is a look-up table that translates *code* into something meaningful. The **code** is a list of numbers corresponding different lines in the file.

Here we use the `W08_words.txt` file as a codebook. The code `41879 62760` corresponds to the words at line 41879 (good) and 62760 (morning) in the file and thus can be translated to `good morning`.

Please complete the code below

```python
def read_codebook(fname = 'W08_words.txt'):    
    '''return a list of words'''
    # INSERT YOUR CODE BELOW
    # Store each line as an elment of a list, then retunr the list
    # Hint, use strip() method to remove '\n'
    # ~ 3 lines
    with open(fname, 'r', encoding='utf8') as fin:
        words = [line.strip() for line in fin]
    return words


def encrypt_words(words, codebook):
    # INSERT YOUR CODE BELOW
    # Look up code for each words and store them into a list li
    # ~ 3 lines
    li = [codebook.index(word.lower()) for word in words.split()]
    return li

def decrypt_codes(code, codebook):
    # INSERT YOUR CODE BELOW
    # Look up words from code, and store them into a string
    # ~ 4 lines
    s = ' '.join([codebook[c] for c in code])
    return s

# TESTCASES BELOW
cb = read_codebook()
li = encrypt_words('good morning', cb)
print(li) #[41879, 62760]

code = [44543, 0, 64852, 24327]
s = decrypt_codes(code, cb)
print(s) # have a nice day
```

Output:
```
[41879, 62760]
have a nice day
```

Tests for Task1 (points: 50)
```
>>> "The Fallohides, the least numerous," in read_page(4)
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> "That name the hobbits only knew in legends" in read_page(57)
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Task2 (points: 50)
```
>>> encrypt_words('happy birthday', cb) == [44168, 9866]
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> 'a blessing in disguise' in decrypt_codes([0, 10268, 48608, 27628], cb)
True
```
failure_message: Incorrect answer, try again!
success_message: Good job!
