# Practical Exam 1 (2023 S1)

**Time:** 110 minutes (105 minutes for solving the tasks + 5 minutes for the upload)<br>
**Mark:** 110 available, but capped at 100.

See full exam information: [here](https://myuni.adelaide.edu.au/courses/85258/pages/prac-exam-1-information?module_item_id=2928787)

<font color='red'>Unless otherwise stated, questions are marked using the autograder. As such, you have to use the function names, parameters, and returns exactly as instructed. Once results come out, you will be able to submit a ticket in Gradescope to get a question(s) reviewed and remarked if necessary.</font><br>

<font color=red>Remember to submit your Jupyter notebook to **Gradescope** before the exam completion. It is recommended you make submissions throughout the exam.</font><br>

<font color=red>Do not discuss or share exam questions anywhere within two weeks after exam completion.</font><br>

---

# In questions 1 ~ 6 you will demonstrate your ability to code the essentials

Total marks 70. The purpose of these questions is to verify you have learnt the **key concepts** of the course.

<!-- BEGIN QUESTION -->

## Q1: User Input & print (5 marks)
Write a program that takes in two numbers of user input as floats. Assign these to the variables `a` and `b`. Then, **prints** the sum as an equation.

**Example (<font color=red>red text</font> represents user input):**
<pre>
Please enter a number: <font color=red>42</font>
Please enter a number: <font color=red>3.3</font>
42.0 + 3.3 = 45.3
</pre>

*This question is manually marked.* 

```python
# INSERT YOUR CODE BELOW (~ 3 lines)
a = input("Please enter a number: ")
a = float(a)
b = input("Please enter a number: ")
b = float(b)
c = a + b
print(f"{a} + {b} = {c}")
```

Output:
```
42.0 + 3.3 = 45.3
```

<!-- END QUESTION -->

## Q2: Functions (15 marks)
Write a function `print_greeting()` that has two parameters `greeting` and `name` (in that order) with default values of 'Hello' and 'World' respectively.

This function should **print** a message of greeting the provided name (see the example test cases below).

```python
# INSERT YOUR CODE BELOW (~ 1 line)
def print_greeting(greeting,name):
    greeting = input("Please enter your greeting: ")
    greeting = str(greeting)
    print(greeting + name)
    print('Hello'+ 'World')
    print('Hello'+ name )
    
# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
print_greeting('Thank you,', 'driver') #'Thank you, driver' -- PUBLIC TEST CASE
print_greeting() #'Hello World'
print_greeting(name='Mr Bennett') #'Hello Mr Bennett'
```

Output:
```
Thank youdriver
HelloWorld
Hellodriver
TypeError                                 Traceback (most recent call last)
Cell In[74], line 13
      9 # =======================
     10 # EXAMPLE TESTING CASES
     11 # DO NOT CHANGE
     12 print_greeting('Thank you,', 'driver') #'Thank you, driver' -- PUBLIC TEST CASE
---> 13 print_greeting() #'Hello World'
     14 print_greeting(name='Mr Bennett') #'Hello Mr Bennett'

TypeError: print_greeting() missing 2 required positional arguments: 'greeting' and 'name'
```

## Q3: Booleans and if-statements (15 marks)
Implement the following function such that it returns `True` or `False` depending on the following rules:
* If `n` is divisible 3 or 4 and `n` is not divisble by 12, the function should return `True`.
* Otherwise, the function returns `False`.

```python
# INSERT YOUR CODE BELOW
def div_by_3_or_4_but_not_12(n):
    # INSERT YOUR CODE BELOW (~ 4-8 lines)
    if n / 3 or n / 4 and n //12:
        return True
    else:
        return False
    
    

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
print(div_by_3_or_4_but_not_12(10)) #False: 10 is not divisible by any of the numbers -- PUBLIC TEST CASE
print(div_by_3_or_4_but_not_12(3))  #True: The remainder of 3/3 is 0
```

Output:
```
True
True
```

<!-- BEGIN QUESTION -->

## Q4: Loops #1 (15 marks)
Using a **while loop**, implement the `spam` function. This function should first ask the user to enter a message. Then, it should ask them to enter a number. Afterwards the function prints the user's message that number of times (each being on a new line).

**Example (<font color=red>red text</font> represents user input):**
<pre>
Enter a message: <font color=red>hello world</font>
Please enter a number: <font color=red>2</font>
hello world
hello world
</pre>

*This question is manually marked.*

```python
def spam(a,b):
    # INSERT YOUR CODE BELOW (~ 6 lines)
    a = input("Enter a message: ")
    a = str(a)
    b = input("Please enter a number: ")
    b = int(b)
    return spam(a,b)
    
# =======================
# this will run the function to allow you to test, 
# spam() # COMMENT THE LINE OUT WHEN YOU SUBMIT to ensure a speedy submission
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Q5: Loops #2 (15 marks)
Using a **for loop**, implement a function that **returns** a string containing every second number from `0` upto  and including `n` (separated by spaces '` `').

*You may assume that n is always a positive integer. You do not need to trim trailing whitespace.*

*This question is manually marked.*

```python
# INSERT YOUR CODE BELOW
def generate_sequence(n):
    text = ''
    # INSERT YOUR CODE BELOW (~ 2 - 5 lines)
    n >= 0
    for text in range (n):
        text == text + 1
        return text

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
print(generate_sequence(10)) #0 2 4 6 8 10 
print(generate_sequence(0))  #0 
print(generate_sequence(4))  #0 2 4 
```

Output:
```
0
None
0
```

<!-- END QUESTION -->

## Q6: List of lists (5 marks)
List `li` is a nested list. Implement the `get_sublist_length` function that **returns** the length of the `n`'th list of `li`.

*You may assume that* `n` *is always a positive integer, and* `li` *contains at least* `n+1` *elements.*

```python
# INSERT YOUR CODE BELOW
def get_sublist_length(li, n):
    # INSERT YOUR CODE BELOW (~ 1-2 lines)
    return len(li[n])

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
li = [[0,1],[2,3],[4,5],[6,7],[8,9]]
print(get_sublist_length(li, 3)) #2 -- PUBLIC TEST CASE
li = [[0],[0,0],[0,0,0]]
print(get_sublist_length(li, 2)) #3
```

Output:
```
2
3
```

---

# In questions 7 ~ 8 you will combine your knowledge of the fundamentals to create useful programs
Total marks 40. **These questions are designed to be more challenging** in order to identify those working at a Distinction level or higher. **As each question is made up of two parts, even if you cannot complete everything you can earn some marks!**

## Q7: Convert a mark to grade (20 marks)
The University of Adelaide uses the following grading scheme for its honours courses.
- 1:   mark between 80 - 100
- 2A:  mark between 70 - 89
- 2B:  mark between 60 - 69
- 3:   mark between 50 - 59
- NAH: mark between 0 - 49

### Part A (10 marks)
Complete the function `mark_to_grade` that converts a `mark` (integer) into a `grade` (string) and **returns** that `grade`. If the `mark` is below 0 or above 100, the function should **return** `'invalid'`.

### Part B (10 marks)
Using `mark_to_grade`, complete the function `list_mark_to_grade` that converts a list of marks into a list of grades.

```python
def mark_to_grade(mark):
    # INSERT YOUR CODE BELOW (~ 12 lines)
    mark = int(mark)
    if mark <= 0 or mark >= 49:
        grade = print("NAH")
    elif mark == 50 or mark >= 59:
        grade = print("3")
    elif mark == 60 or mark >= 69:
        grade = print("2B")
    elif mark == 70 or mark >= 79:
        grade = print("2A")
    elif mark == 80 or mark >= 100:
        grade = print("1")
    elif mark < 0 or mark > 100:
        grade = print("invalid")
    grade = str(grade)

def list_mark_to_grade(marks):
    output = []
    # INSERT YOUR CODE BELOW (~ 2 - 5 lines)
    print(li[marks])
    return output

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
print(mark_to_grade(70)) #'2A'
print(mark_to_grade(-5)) #'invalid'
print(list_mark_to_grade([10, 50, 60, 70])) #['NAH', '3', '2B', '2A']
print(list_mark_to_grade([30, 99, -5, 70])) #['NAH', '1', 'invalid', '2A']
```

Output:
```
NAH
None
NAH
None
Cell In[69], line 29
     27 print(mark_to_grade(70)) #'2A'
     28 print(mark_to_grade(-5)) #'invalid'
---> 29 print(list_mark_to_grade([10, 50, 60, 70])) #['NAH', '3', '2B', '2A']
     30 print(list_mark_to_grade([30, 99, -5, 70])) #['NAH', '1', 'invalid', '2A']

Cell In[69], line 21, in list_mark_to_grade(marks)
     19 output = []
     20 # INSERT YOUR CODE BELOW (~ 2 - 5 lines)
---> 21 print(li[marks])
     22 return output

TypeError: list indices must be integers or slices, not list
```

## Q8: Critical Hit!
Your friend loves playing a tabletop roleplaying games and they have asked you to create a program to automate some of the experience.

In the game, performing an *attack* involves two sets of rolls. A roll to hit and a roll for calculating damage.

### Part A (10 marks)
Implement the `roll_to_hit(armour, mod)` function.

First simulate rolling a 20-sided die using the `random` module and assign it to a variable `d20`. Then, **return** one of the following integers:
* return `2` if `d20` is 20,
* return `0` if `d20` is 1 OR `d20` plus `mod` is strictly less than `armour`, or
* return `1` in any other case

### Part B (10 marks)
Implement the `calculate_damage(numDice, numSided, mod)` function. The parameters are as follows:
* `numDice` is the number of dice rolls to simulate
* `numSided` is the number of sides on the dice you are simulating (not that all the dice you roll are the same type)
* `mod` an integer that is added to the returned result **only if** the `numDice` is greater than 0

Your function should **return** the amount of damage caused by an attack.
* return `0`, if the `numDice` is `0` (regardless of `mod`).
* otherwise, return damage using the following formula: $damage = mod + roll\_sum$, where `roll_sum` is defined as the combined value of rolling the `numSided` die `numDice` times.
  * For example, if `numDice` was 2, it would return damage as $damage = mod + roll\_1's\_result + roll\_2's\_result$

<font color=red>**DO NOT CHANGE THE SEED OR IMPORTED MODULES**</font>

```python
import random # DO NOT change this line
random.seed(5) # DO NOT change this line

def roll_to_hit(armour, mod):
    # INSERT YOUR CODE BELOW (1 lines)
    d20 = random(armour, mod)
    
    # INSERT YOUR CODE BELOW (~6 lines)
    ...
    
def calculate_damage(numDice, numSided, mod):
    roll_sum = 0
    # INSERT YOUR CODE BELOW (~5 lines)
    ...

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
e1 = roll_to_hit(4, 0)  #2 (rolled a 20 which always returns 2)
e2 = roll_to_hit(20, 0) #0 (rolled a  9+0 which is less than 20)
e3 = roll_to_hit(8, 1)  #1 (rolled a 12+1 which is greater than 8)
print(f'{e1}, {e2}, {e3}') #2, 0, 1 (if you get these three numbers, move onto Part B)

print(f'rolled {e1*1}d4 + 0 and dealt {calculate_damage(e1*1, 4, 0)} damage') # (1 + 4) + 0 = 5 damage
print(f'rolled {e2*1}d8 + 2 to deal {calculate_damage(e2*1, 8, 2)} damage')   # rolling 0 dice means no damage
print(f'rolled {e3*3}d10 + 1 to deal {calculate_damage(e3*3, 10, 1)} damage') # (4+1+3) + 0 = 9 damage
```

---

# <font color=red>Remember to submit your Jupyter notebook to Gradescope</font>

Tests for Q1 (points: 10)


Tests for Q2 (points: 15)
```
>>> print_greeting('Thank you,', 'driver') #'Thank you, driver'
Thank you, driver
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Q3 (points: 15)
```
>>> div_by_3_or_4_but_not_12(10) #False: 10 is not divisible by any of the numbers -- PUBLIC TEST
False
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Q4 (points: 20)


Tests for Q6 (points: 5)
```
>>> li = [[0,1],[2,3],[4,5],[6,7],[8,9]]
>>> get_sublist_length(li, 3) #2 -- PUBLIC TEST CASE
2
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Q7 (points: 20)


Tests for Q8 (points: 20)
