# Coding Exam 2 (2023 S1)

**Time:** 110 minutes<br>
**Maximum mark:** 100

<font color='red'>Unless otherwise stated, questions are marked using the autograder. As such, you have to use the function names, parameters, and returns exactly as instructed. Once results come out, you will be able to submit a ticket in Gradescope to get a question(s) reviewed and remarked if necessary.</font><br>

<font color=red>Remember to submit your Jupyter notebook to **Gradescope** before the exam completion. It is recommended you make submissions throughout the exam.</font><br>

<font color=red>Do not discuss or share exam questions anywhere within two weeks after exam completion.</font><br>

---

# In questions 1 ~ 6 you will demonstrate your ability to code the essentials

Total marks 70.

<!-- BEGIN QUESTION -->

## Q1 : User input & print (5 marks)

Write a program that takes in user input string followed by an integer. Assign these to the variables `m` and `n` and print the string `m` `n`-times **without using a loop**. There should not be any spaces of new lines between each instance of `m`.

**Example (<font color=red>red text</font> represents user input):**
<pre>
Please enter a string: <font color=red>la</font>
Please enter a number: <font color=red>3</font>
lalala
</pre>

*This question is manually marked. You may assume that the user always assigns n to a positive integer.* 

```python
# YOUR CODE BELOW
s = input("Please enter a string: ")
n = int(input("Please enter a number: "))

count = 0
while count < n:
    print(s, end="")
    count += 1

```

Output:
```
lalala
```

<!-- END QUESTION -->

## Q2 : Functions, conditionals & loops (10 marks)

Write a function `occurances()` with 2 parameters `string` and `letter`. This function should **return** the number of times `letter` appeared in `string`.

You may assume `string` and `letter` only use lowercase characters and spaces.

```python
# INSERT YOUR CODE BELOW
def occurances(string, letter):
    count = 0
    for char in string:
        if char == letter:
            count += 1
    return count

# =======================
# EXAMPLE TESTING CASES
print(occurances('hello world', 'l')) #3 -- PUBLIC TEST CASE
print(occurances('zzz', 'a')) #0
```

Output:
```
3
0
```

## Q3: Lists & loops (15 marks)

Write the function `print_elements` that has 2 parameters.
* The first `li` parameter is list.
* The second parameter `n` is an integer. This parameter is optional and if it is not provided, its value defaults to 1 and the function behaves as normal.

The function must **return** a new list containing every `n`'th value of the list as a new list (starting from the 0th element).

*You can assume `li` will always be the first parameter with at least 1 element.*

```python
def print_elements(li, n=1):
    return li[::n]
    
# =======================
# EXAMPLE TESTING CASES
li = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
print(print_elements(li, 1)) #[0, 1, 2, 3, 4, 5, 6, 7, 8, 9] -- PUBLIC TEST CASE
print(print_elements(li))    #[0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
print(print_elements(li, 3)) #[0, 3, 6, 9]
```

Output:
```
[0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
[0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
[0, 3, 6, 9]
```

## Q4: Files & exceptions (10 marks)

Complete the `reverse_char` function which reverses the content of `file1` character by character, then saves the reversed content to `file2`. You must also handle the exception caused by `file1` not existing. In this case, you must **print** "X does not exist", where X is `file1`.

**Ensure you download the appropriate files from the Exam Page on MyUni. Do not change the content of these files.**

```python
def reverse_char(file1, file2):
    try:
        with open(file1, 'r') as i:
            content = i.read()
            reversed_content = content[::-1]
        
        with open(file2, 'w') as i:
            i.write(reversed_content)
    except FileNotFoundError:
        print(file1, "does not exist")



# =======================
# EXAMPLE TESTING CASES
reverse_char('practical_exam_input0.txt', 'out.txt')
# file out.txt should contains:
#
# ;htworgrednu eht ni tneb ti erehw oT
# dluoc I sa raf sa eno nwod dekool dnA
# doots I gnol ,relevart eno eb dnA
# htob levart ton dluoc I yrros dnA
# ,doow wolley a ni degrevid sdaor owT
```

## Q5 : Dictionaries (15 marks)

Complete the following functions:
- `build_dictionary()` that creates and **return** a dictionary that has 26 keys, one for each letter in the English alphabet. The value of 'a' is 26, 'b' is 25, etc. all the way to 'z' is 1.
- `lazy_hash(m, d)` that takes in message `m` and dictionary `d`. The function **returns** the total value of all characters in `m`. Spaces and symbols have no impact on the value. In the event of an empty string, return 0.

**Do not change any variable names or function signatures of this question**

```python
import string
# print(string.ascii_lowercase) # hint: this line of code gives you lowercase alphabets a to z.

def build_dictionary():
    # INSERT YOUR CODE BELOW
    dictionary = {}
    for i, char in enumerate(range(ord('a'), ord('z') + 1)):
        dictionary[chr(char)] = 26 - i
    return dictionary


def lazy_hash(m, d):
    # INSERT YOUR CODE BELOW
    if not m:
        return 0

    total_value = 0
    for char in m:
        if char.isalpha():
            char = char.lower()
            if char in d:
                total_value += d[char]
    return total_value


# =======================
# EXAMPLE TESTING CASES
my_dict = build_dictionary()
print(my_dict)
print(lazy_hash('abc', my_dict)) # 75 (26 + 25 + 24 = 75)
print(lazy_hash('a z', my_dict)) # 27 (26 + 1 = 27)
print(lazy_hash('b-y', my_dict)) # 27
print(lazy_hash('ooo', my_dict)) # 36
```

Output:
```
{'a': 26, 'b': 25, 'c': 24, 'd': 23, 'e': 22, 'f': 21, 'g': 20, 'h': 19, 'i': 18, 'j': 17, 'k': 16, 'l': 15, 'm': 14, 'n': 13, 'o': 12, 'p': 11, 'q': 10, 'r': 9, 's': 8, 't': 7, 'u': 6, 'v': 5, 'w': 4, 'x': 3, 'y': 2, 'z': 1}
75
27
27
36
```

## Q6 : Mutable Data, lists & tuples (15 marks)

Implement the two following functions, each of which take in a list `li` and integer `n`.

- `transformIntoMinimumList()`: Modifies the original list so that it contains only `n` smallest values.
  - Hint: You can use `li.clear()` to clear a list with name `li`.
- `getSecondMinMax()`: which takes in a list `li` as an argument and **returns** a tuple of the second smallest and second largest values, in that order. You should not modify the original list `li`.

For simplicity, you can assume that `li` only contains integers, has at least `n` items and no fewer than 4 items.

The following functions may be useful.
- You can use `li.clear()` to clear a list with name `li`.
- `min()` and `max()`, see [Python Documentation](https://docs.python.org/3/library/functions.html)

```python
def transformIntoMinimumList(li, n):
    for i in range (0, n+1):
     li.clear()
    return n
        
def getSecondMinMax(li):
   if len(li) < n:
      return li
    
    
# =======================
# EXAMPLE TESTING CASES
li = [1, 2, 3, 4, -1, 5, 6]
transformIntoMinimumList(li, 2)
print(li) # [-1, 1]

li = [1, 2, 3, 4, -1, 5, 6]
print(getSecondMinMax(li)) # (1, 5)
```

Output:
```
[]
None
```

# In questions 7 ~ 8 you will combine your knowledge of the fundamentals to create useful programs
Total marks 40. These questions are designed to be more challenging. Each question is made up of two parts, so if you cannot get everything implement as many of the features as you can!

## Q7: Hash Tables (20 marks)

Complete `hashTable()` function that takes string `msg` and Boolean `letter_mode`. 

- `msg` is a string of completely lowercase words seperated by spaces.
- `letter_mode` is a Boolean that determines what the keys of the dictionary are.
  - if True, the keys are the first letter of the words.
  - if False, the keys are the lengths of the words.

The function **returns** a dictionary of lists. The dictionary is made of key value pairs. The values in the dictionary are the words of string `msg` with keys as defined above.

Do **not** sort the keys or values, simply add them in the order you encounter them.

Examples:
1) `hashTable('hello nice ninja', True)` should return `{'h': ['hello'], 'n': ['nice', 'ninja']}`. as only 'hello' begins with 'h', while both 'nice' and 'ninja' begin with 'n'.
2) `hashTable('abra kadabra alakazam', False)` should return `{4: ['abra'], 7: ['kadabra'], 8: ['alakazam']}` as both 'abra' and 'kadabra' end with 'a' and only 'alakazam' ends with 'm'.

*10 marks test for each of the two Boolean values.*

```python
def hashTable(msg, letter_mode):
    return (msg, letter_mode)
    
# =======================
# EXAMPLE TESTING CASES
print(hashTable('hello nice ninja', True)) # {'h': ['hello'], 'n': ['nice', 'ninja']}
print(hashTable('abra kadabra alakazam', False)) # {4: ['abra'], 7: ['kadabra'], 8: ['alakazam']}
```

Output:
```
('hello nice ninja', True)
('abra kadabra alakazam', False)
```

## Q8: Calculating Statistics (20 marks)

Complete each of the following functions that have a list of integers `numbers` as their parameter. For each question, you can assume that `numbers` will have at least 2 elements, and all elements are valid integers. If you cannot complete a part, try one of the other functions.

### Part A (3 marks)
Implement the function `calculate_mean`. This function should **return** the average of the numbers in the list (i.e. the sum of all numbers in the list divided by how many numbers are in the list). **Do not round this number.**

### Part B (6 marks)
Implement the function `calculate_median`. This function should **return** the middle value of the sorted list. If the list has an even number of elements, take the average of the two middle values **(in which case do not round).**

**Hint:** the function `sorted()` might be useful; see [Python Documentation](https://docs.python.org/3/library/functions.html)

### Part C (9 marks)
Write a Python function called `calculate_mode` that **returns** a list of the most frequently occurring value(s) in the list.

**Hint:** Consider using a dictionary with keys representing each unique number in the list and values corresponding to the number of times that number appeared.

### Part D (2 marks)
Write a Python function called `calculate_statistics()` that **returns** a formatted string containing of the following statistics that matches the formatting demonstrated in the example below:
- "mean": from part A
- "median": from part B
- "mode": from part C

Example:
- Given an input of `[2, 4, 6, 4, 8, 2, 10, 2, 4]`, its output should be "`mean = 4.666666666666667, median = 4, mode = [2, 4]`" (without quotes).
- Given an input of `[2, 4, 6, 8, 2, 14]`, its output should be "`mean = 6.0, median = 5.0, mode = [2]`" (without quotes).

```python
def calculate_mean(numbers):
    return sum(numbers) / len(numbers)

def calculate_median(numbers):
    sorted_num = sorted(numbers)
    n = len(sorted_num)
    if n % 2 == 0:
        mid1 = sorted_num[n // 2 - 1]
        mid2 = sorted_num[n // 2]
        return (mid1 + mid2) / 2
    else:
        return sorted_num[n // 2]

def calculate_mode(numbers):
    frequency = {}
    max_frequency = 0
    modes = []
    for num in numbers:
        if num not in frequency:
            frequency[num] = 1
        else:
            frequency[num] += 1
        if frequency[num] > max_frequency:
            max_frequency = frequency[num]
    for num, count in frequency.items():
        if count == max_frequency:
            modes.append(num)
    return modes


def calculate_statistics(numbers):
    mean = calculate_mean(numbers)
    median = calculate_median(numbers)
    mode = calculate_mode(numbers)
    return f"mean = {mean}, median = {median}, mode = {mode}"

# =======================
# EXAMPLE TESTING CASES
print("## TEST #1")
numbers = [2, 4, 6, 4, 8, 2, 10, 2, 4]

r1 = calculate_mean(numbers)
print(r1) # 4.666666666666667

r2 = calculate_median(numbers)
print(r2) # 4

r3 = calculate_mode(numbers)
print(r3) # [2, 4]

r4 = calculate_statistics(numbers)
print(r4) # "mean = 4.666666666666667, median = 4, mode = [2, 4]"

print("\n## TEST #2")
numbers = [2, 4, 6, 8, 2, 14]

r1 = calculate_mean(numbers)
print(r1) # 6.0

r2 = calculate_median(numbers)
print(r2) # 5.0

r3 = calculate_mode(numbers)
print(r3) # [2]

r4 = calculate_statistics(numbers)
print(r4) # "mean = 6.0, median = 5.0, mode = [2]"
```

Output:
```
## TEST #1
4.666666666666667
4
[2, 4]
mean = 4.666666666666667, median = 4, mode = [2, 4]

## TEST #2
6.0
5.0
[2]
mean = 6.0, median = 5.0, mode = [2]
```

# <font color=red>Remember to submit your Jupyter notebook on Gradescope</font>

Tests for Q2 (points: 10)
```
>>> occurances('hello world', 'l') #3 -- PUBLIC TEST CASE
3
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Q3 (points: 15)
```
>>> l1 = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
>>> l2 = print_elements(l1, 1) # PUBLIC TEST CASE
>>> 
>>> assert l1 == [0, 1, 2, 3, 4, 5, 6, 7, 8, 9], "original list was modified"
>>> assert l2 == [0, 1, 2, 3, 4, 5, 6, 7, 8, 9], "final list is incorrect"
>>> 
>>> True
True
```


Tests for Q4 (points: 10)


Tests for Q5 (points: 15)


Tests for Q6 (points: 15)


Tests for Q7 (points: 20)


Tests for Q8 (points: 20)
```
>>> numbers = [2, 4, 6, 4, 8, 2, 10, 2, 4]
>>> calculate_mean(numbers) # 4.666666666666667
4.666666666666667
```
```
>>> numbers = [2, 4, 6, 4, 8, 2, 10, 2, 4]
>>> calculate_median(numbers) # 4
4
```
```
>>> numbers = [2, 4, 6, 4, 8, 2, 10, 2, 4]
>>> calculate_mode(numbers) # [2, 4]
[2, 4]
```
